# Tentamen template

* We give examples of the diferent types of questions anf how they are graded
* For manual grading it is easiest to start with all manually graded questions, because when you do manual grading you grade by notebook from top to bottom and the manually graded cells are not easily recognizable.

```python
# Run this cell first

import pandas as pd
import numpy as np
from nose.tools import assert_equal

observed= np.array([1,3,5])
predicted= np.array([1.5,2,6])
```

# Manually graded example

* Answer in MarkDown

**Q3** Give the formula for  the _root mean squared error_ between the two lists of samples: `observed` and `predicted`.
    
Use Latex.    

<!-- grade_id: vraag-2, points: 2, grade: true, solution: true -->
 
$$ \sqrt{\frac{\Sigma_1^N (\mathtt{observed}  -\mathtt{predicted})^2}{N}}$$

# Manually graded example

* Answer in code

**Q4** Give the formula for  the _root mean squared error_ between the two lists of samples: `observed` and `predicted`.
    


<!-- grade_id: vraag-3, points: 1, grade: true, solution: true -->
```python
rmse = None
rmse
### BEGIN SOLUTION
np.sqrt(((observed-predicted)**2).mean()),(((observed-predicted)**2).mean())**.5
### END SOLUTION
```

Output:
```
(0.8660254037844386, 0.8660254037844386)
```

# Autograded example

**Q1** Compute the _root mean squared error_ between the two lists of samples: `observed` and `predicted`.

* First you write a function for that. Do that in the next cell.

<!-- grade_id: vraag-1-a, points: 0, grade: false, solution: true -->
```python
def rmse(observed,predicted):
    '''Compute the root mean squared error between the two lists of samples: observed and predicted.'''
    ### BEGIN SOLUTION
    return np.sqrt(((observed-predicted)**2).mean())
    ### END SOLUTION

# test 
rmse(observed,predicted)
```

Output:
```
0.8660254037844386
```

<!-- grade_id: vraag-1-a-test, points: 1, grade: true, solution: false -->
```python
# Hier staat een onzichtbare test
### BEGIN HIDDEN TESTS

assert_equal(rmse(observed,predicted), np.sqrt(((observed-predicted)**2).mean()), 'Your result does not compute np.sqrt(((observed-predicted)**2).mean()')
### END HIDDEN TESTS
```

**Q2** Now again you compute the RMSE between the two lists of samples: `observed` and `predicted`, but you store the result in the variable `Rmse` in the cell below.

<!-- grade_id: vraag-1, points: 0, grade: false, solution: true -->
```python
Rmse = None # replace with your answer
Rmse
### BEGIN SOLUTION
Rmse = np.sqrt(((observed-predicted)**2).mean())
### END SOLUTION
```

<!-- grade_id: vraag-1-test, points: 1, grade: true, solution: false -->
```python
# Hier staat een onzichtbare test
### BEGIN HIDDEN TESTS

assert_equal( Rmse, np.sqrt(((observed-predicted)**2).mean()), 'Your result does not compute np.sqrt(((observed-predicted)**2).mean()')
### END HIDDEN TESTS
```

# Autograded example 2

* Here the answer is simply a fixed value.
* We do not need to check code.

**Q2a** Compute the RMSE  between the two lists of samples `observed` and `predicted` given above.

* You simply give the answer, rounded to 2 digits.
* You do not give code. (So your answer has this shape: `answer= 7.76` )

 

<!-- grade_id: vraag-2a, points: 0, grade: false, solution: true -->
```python
answer= None # replace with your answer
answer
### BEGIN SOLUTION
 
answer =round(np.sqrt(((observed-predicted)**2).mean()),2)
### END SOLUTION
```

<!-- grade_id: vraag_2a-test, points: 1, grade: true, solution: false -->
```python
# Hier staat een onzichtbare test
### BEGIN HIDDEN TESTS

assert_equal( answer, 0.87, 'Your result does not compute np.sqrt(((observed-predicted)**2).mean()')
### END HIDDEN TESTS
```