```python
import pandas as pd
import numpy as np
```

## Welkom bij het tentamen

### 1. Schrijf een functie

<!-- grade_id: vraag-1, points: 0, grade: false, solution: true -->
```python
# Dit en dat moet je doen

def functie(s):
    return s * 2
```

<!-- grade_id: vraag-1-test, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

assert_equal(functie('test'), 'testtest', 'Functie doet niet wat hij moet doen')
### END HIDDEN TESTS
```

## 2. Leg uit ....

<!-- grade_id: vraag-2, points: 3, grade: true, solution: true -->
Bla die bla ....

## 3. Doe nog iets

<!-- grade_id: vraag-3, points: 1, grade: true, solution: true -->
```python
antwoord = 5.5

# JOUW CODE HIER
```

```python

```